The coarse mental map of the breast is anchored on the nipple

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1 (Public review):

The statistically adequate way of testing the biases is a hierarchical regression model (LMM) with a distance of the physical location from the nipple as a predictor, and a distance of the reported location from the nipple as a dependent variable. Either variable can be unsigned or signed for greater power, for example, coding the lateral breast as negative and the medial breast as positive. The bias will show in regression coefficients smaller than 1.

Thank you for this suggestion. We have subsequently replaced the relevant ANOVA analyses with LMM analyses. Specifically, we use an LMM for breast and back separately to show the different effects of distance, then use a combined LMM to compare the interaction. Finally, we use an LMM to assess the differences between precision and bias on the back and breast. The new analysis confirms earlier statements and do not change the results/interpretation of the data.

Moreover, any bias towards the nipple could simply be another instance of regression to the mean of the stimulus distribution, given that the tested locations were centered on the nipple. This confound can only be experimentally solved by shifting the distribution of the tested locations. Finally, given that participants indicated the locations on a 3D model of the body part, further experimentation would be required to determine whether there is a perceptual bias towards the nipple or whether the authors merely find a response bias.

A localization bias toward the nipple in this context does not show that the nipple is the anchor of the breast's tactile coordinate system. The result might simply be an instance of regression to the mean of the stimulus distribution (also known as experimental prior). To convincingly show localization biases towards the nipple, the tested locations should be centered at another location on the breast.

Another problem is the visual salience of the nipple, even though Blender models were uniformly grey. With this type of direct localization, it is very difficult to distinguish perceptual from response biases even if the regression to the mean problem is solved. There are two solutions to this problem: 1) Varying the uncertainty of the tactile spatial information, for example, by using a pen that exerts lighter pressure. A perceptual bias should be stronger for more uncertain sensory information; a response bias should be the same across conditions. 2) Measure bias with a 2IFC procedure by taking advantage of the fact that sensory information is noisier if the test is presented before the standard.

We believe that the fact that we explicitly tested two locations with equally distributed test locations, both of which had landmarks, makes this unlikely. Indeed, testing on the back is exactly what the reviewer suggests. It would also be impossible to test this “on another location on the breast” as we are sampling across the whole breast. Moreover, as markers persisted on the model within each block, the participants were generating additional landmarks on each trial. Thus, if there were any regression to the mean, this would be observed for both locations. Nevertheless, we recognize that this test cannot distinguish between a sensory bias towards the nipple and consistent response bias that is always in the direction of the nipple, though to what extent these are the same thing is difficult to disentangle. That said, if we had restricted testing to half of the breast such that the distribution of points was asymmetrical this would allow us to test the hypothesis put forward by the reviewer. We recognize that this is a limitation of the data and have downplayed statements and added caveats accordingly.

We have changed the appropriate heading and text in the discussion to downplay the finding:

“Reports are biased towards the nipple”

“suggesting that the nipple plays a pivotal role in the mental representation of the breast.”

it might be harder to learn the range of locations on the back given that stimulation is not restricted to an anatomically defined region as it is the case for the breast.

We apologize for any confusion but the point distribution is identical between tasks, as described in the methods.

The stability of the JND differences between body parts across subjects is already captured in the analysis of the JNDs; the ANOVA and the post-hoc testing would not be significant if the order were not relatively stable across participants. Thus, it is unclear why this is being evaluated again with reduced power due to improper statistics.

We apologize for any confusion here. Only one ANOVA with post-hoc testing was performed on the data. The second parenthetical describing the test was perhaps redundant and confusing, so I have removed it.

“(Error! Reference source not found.A, B, 1-way ANOVA with Tukey’s HSD post-hoc t-test: p = 0.0284)”

The null hypothesis of an ANOVA is that at least one of the mean values is different from the others; adding participants as a factor does not provide evidence for similarity.

We agree with this statement and have removed the appropriate text.

The pairwise correlations between body parts seem to be exploratory in nature. Like all exploratory analyses, the question arises of how much potential extra insights outweigh the risk of false positives. It would be hard to generate data with significant differences between several conditions and not find any correlations between pairs of conditions. Thus, the a priori chance of finding a significant correlation is much higher than what a correction accounts for.

We broadly agree with this statement. However, we believe that the analyses were important to determine if participants were systematically more or less acute across body parts. Moreover, both the fact that we actually did not observe any other significant relationships and that we performed post-hoc correction imply that no false positives were observed. Indeed, in the one relationship that was observed, we would need to have an assumed FDR over 10x higher than the existing post hoc correction required implying a true relationship.

If the JND at mid breast (measured with locations centered at the nipple) is roughly the same size as the nipple, it is not surprising that participants have difficulty with the categorical localization task on the nipple but perform better than chance on the significantly larger areola.

We agree that it is not surprising given the previously shown data, however, the initial finding is surprising to many and this experiment serves to reinforce the previous finding.

Neither signed nor absolute localization error can be compared to the results of the previous experiments. The JND should be roughly proportional to the variance of the errors.

We apologize for any confusion, however we are not comparing the values, merely observing that the results are consistent.

Reviewer #2 (Public review):

I had a hard time understanding some parts of the report. What is meant by "broadly no relationship" in line 137?

We have removed the qualifier to simplify the text.

It is suggested that spatial expansion (which is correlated with body part size) is related between medial breast and hand - is this to say that women with large hands have large medial breast size? Nipple size was measured, but hand size was not measured, is this correct?

Correct. We have added text to state as such.

It is furthermore unclear how the authors differentiate medial breast and NAC. The sentence in lines 140-141 seems to imply the two terms are considered the same, as a conclusion about NAC is drawn from a result about the medial breast. This requires clarification.

Thank you for catching this, we have corrected it in the text.

Finally, given that the authors suspect that overall localization ability (or attention) may be overshadowed by a size effect, would not an analysis be adequate that integrates both, e.g. a regression with multiple predictors?

If the reviewer means that participants would be consistently “acute” then we believe that SF1 would have stronger correlations. Consequently, we see no reason to add “overall tactile acuity” as a predictor.

In the paragraph about testing quadrants of the nipple, it is stated that only 3 of 10 participants barely outperformed chance with a p < 0.01. It is unclear how a significant ttest is an indication of "barely above chance".

We have adjusted the text to clarify our meaning.

“On the nipple, however, participants were consistently worse at locating stimuli on the nipple than the breast (paired t-test, t = 3.42, p < 0.01) where only 3 of the 10 participants outperformed chance, though the group as a whole outperformed chance (Error! Reference source not found.B, 36% ± 13%; Z = 5.5, p < 0.01).”

The final part of the paragraph on nipple quadrants (starting line 176) explains that there was a trend (4 of 10 participants) for lower tactile acuity being related to the inability to differentiate quadrants. It seems to me that such a result would not be expected: The stated hypothesis is that all participants have the same number of tactile sensors in their nipple and areola, independent of NAC size. In this section, participants determine the quadrant of a single touch. Theoretically, all participants should be equally able to perform this task, because they all have the same number of receptors in each quadrant of nipple and areola. Thus, the result in Figure 2C is curious.

We agree that this result seemingly contradicts observations from the previous experiment, however we believe that it relates to the distinction between the ability to perform relative distinctions and absolute localizations. In the first experiment, the presentation of two sequential points provides an implicit reference whereas in the quadrant task there is no reference. With the results of the third experiment in mind, biases towards the nipple would effectively reduce the ability of participants to identify the quadrant. What this result may imply is that the degree of bias is greater for women with greater expansion. We have added text to the discussion to lay this out.

“This negative trend implicitly contradicts the previous result where one might expect equal performance regardless of size as the location of the stimuli was scaled to the size of the nipple and areola. However, given the absence of a reference point, systematic biases are more likely to occur and thus may reflect a relationship between localization bias and breast size.”

This section reports an Anova (line 193/194) with a factor "participant". This doesn't appear sensible. Please clarify. The factor distance is also unclear; is this a categorical or a continuous variable? Line 400 implies a 6-level factor, but Anovas and their factors, respectively, are not described in methods (nor are any of the other statistical approaches).

We believe this comment has been addressed above with our replacement of the ANOVA with an LMM. We have also added descriptions of the analysis throughout the methods.

The analysis on imprecision using mean pairwise error (line 199) is unclear: does pairwise refer to x/y or to touch vs. center of the nipple?

We have clarified this to now read:

“To measure the imprecision, we computed the mean pairwise distance between each of the reported locations for a given stimulus location and the mean reported location.”

p8, upper text, what is meant by "relative over-representation of the depth axis"? Does this refer to the breast having depth but the equivalent area on the back not having depth? What are the horizontal planes (probably meant to be singular?) - do you simply mean that depth was ignored for the calculation of errors? This seems to be implied in Figure 3AB.

This is indeed what we meant. We have attempted to clarify in the text.

“Importantly, given the relative over-representation of the depth axis for the breast, we only considered angles in the horizontal planes such that the shape of the breast did not influence the results.” Became:

“Importantly, because the back is a relatively flat surface in comparison to the breast, errors were only computed in the horizontal plane and depth was excluded when computing the angular error.”

Lines 232-241, I cannot follow the conclusions drawn here. First, it is not clear to a reader what the aim of the presented analyses is: what are you looking for when you analyze the vectors? Second, "vector strength" should be briefly explained in the main text. Third, it is not clear how the final conclusion is drawn. If there is a bias of all locations towards the nipple, then a point closer to the nipple cannot exhibit a large bias, because the nipple is close-by. Therefore, one would expect that points close to the nipple exhibit smaller errors, but this would not imply higher acuity - just less space for localizing anything. The higher acuity conclusion is at odds with the remaining results, isn't it: acuity is low on the outer breast, but even lower at the NAC, so why would it be high in between the two?

Thank you for pointing out the circular logic. We have replaced this sentence with a more accurate statement.

“Given these findings, we conclude that the breast has lower tactile acuity than the hand and is instead comparable to the back. Moreover, localization of tactile events to both the back and breast are inaccurate but localizations to the breast are consistently biased towards the nipple.”

The discussion makes some concrete suggestions for sensors in implants (line 283). It is not clear how the stated numbers were computed. Also, why should 4 sensors nipple quadrants receive individual sensors if the result here was that participants cannot distinguish these quadrants?

Thank you for catching this, it should have been 4 sensors for the NAC, not just the nipple. We have fixed this in the text.

I would find it interesting to know whether participants with small breast measurement delta had breast acuity comparable to the back. Alternatively, it would be interesting to know whether breast and back acuity are comparable in men. Such a result would imply that the torso has uniform acuity overall, but any spatial extension of the breast is unaccounted for. The lowest single participant data points in Figure 1B appear similar, which might support this idea.

We agree that this is an interesting question and as you point out, the data does indicate that in cases of minimal expansion acuity may be constant on the torso. However, in the comparison of the JNDs, post-hoc testing revealed no significant difference between the back and either breast region. Consequently, subsampling the group would result in the same result. We have added a sentence to the discussion stating this.

“Consequently, the acuity of the breast is likely determined initially by torso acuity and then any expansion.”

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Touch plays a key role in our perception of our body and shapes our interactions with the world, from the objects we manipulate to the people we touch. While the tactile sensibility of the hand has been extensively characterized, much less is known about touch on other parts of the body. Despite the important role of the breast in lactation as well as in affective and sexual touch, relatively little is known about its sensory properties. To fill this gap, we investigated the spatial acuity of the breast and compared it to that of the hand and back, body regions that span the range of tactile spatial acuity. First, we found that the tactile acuity of the breast was even lower than that of the back, heretofore the paragon of poor acuity. Second, acuity was lower for larger breasts, consistent with the hypothesis that innervation capacity does not scale with body size. Third, touches to different regions of the nipple were largely indistinguishable, suggesting that the nipple is a sensory unit. Fourth, localization errors were systematically biased toward the nipple and more so at greater distances from the breast.

Revised version: Touch plays a key role in our perception of our body and shapes our interactions with the world, from the objects we manipulate to the people we touch. While the tactile sensibility of the hand has been extensively characterized, much less is known about touch on other parts of the body. Despite the important role of the breast in lactation as well as in affective and sexual touch, relatively little is known about its sensory properties. To fill this gap, we investigated the spatial acuity of the breast and compared it to that of the hand and back, body regions that span the range of tactile spatial acuity. First, we found that the tactile acuity of the breast was even lower than that of the back, heretofore the paragon of poor acuity. Second, acuity was lower for larger breasts, consistent with the hypothesis that innervation capacity does not scale with body size. Third, touches to different regions of the nipple were largely indistinguishable, suggesting that the nipple is a sensory unit. Fourth, localization errors were systematically biased toward the nipple.

Introduction

The sense of touch fulfills a variety of different functions in everyday life, from guiding our interactions with the environment to supporting affective communication and sexual function (1). One of the key properties of touch sensations is that they are localized to a specific part of the body: contact on the shoulder produces a sensation experienced on the shoulder, for example. The precision with which we can localize events on the skin has been shown to be determined by the innervation density at that skin location (2, 3). Because the skin of the fingertips and lips is the most densely innervated, the acuity of these body regions is highest, conferring to us an enhanced ability to distinguish touches on the fingers and lips even when they are close to each other. Innervation density and function are intrinsically related: the fingertips are densely innervated because they account for the vast majority of contacts with objects, and precise information about object interactions is critical to dexterous manipulation (4). In contrast, the acuity on the back is low because precise localization of a touch on the back is of limited value. Skin surface area also plays a role in innervation density; for example, people with large hands exhibit lower acuity than those with small ones (5, 6). This phenomenon is hypothesized to reflect the fact that the number of nerve fibers does not scale with body size: a large body will be more sparsely innervated than a small one given a fixed number of nerve fibers.

While tactile acuity has been extensively studied on the limbs and face, acuity on the torso has received far less experimental attention (7, 8), with the breast being largely ignored beyond its anatomy (9, 10). To fill this gap, we sought to characterize the spatial acuity of the female breast which has roles in lactation, affective touch, and sex. These functions set it apart from other regions of the body. In previous studies, the tactile acuity of the breast was found to be comparably low to that of the back and calf (8). However, the experimental approaches have since received scrutiny to their susceptibility to inaccuracies (11); only the outer breast was tested (excluding the nipple-areolar complex, NAC), and the relationship with breast size was not assessed, or only examined the NAC but not the outer breast (9). Given that breasts grow to vastly different sizes relatively late in life after the nervous system is nearly fully developed (12), the female breast offers a powerful test of the fixed innervation hypothesis, which would predict that people with larger breasts would have poorer spatial acuity than people with smaller ones.

In the present study, we first measured the tactile acuity of two regions of the breast in women– the outer (lateral) breast and medial breast, which includes the NAC and compared these to the hand and back. Second, we examined the relationship between the acuity of the outer breast and breast size. Third, we examined women’s ability to judge the absolute location of touches to their breast. We first found that the spatial acuity of the breast is very low, with similar sensitivity to the back. Second, spatial acuity is inversely correlated with breast size, as predicted from the fixed innervation hypothesis. Third, touches to different parts of the nipple are indistinguishable, suggesting that the nipple is a sensory unit. Fourth, touches on the outer breast are systematically mis-localized as being closer to the nipple than they actually are.

Results

The breast has low spatial acuity

We first measured each participant’s ability to judge the relative position of two touches applied in succession at two nearby locations on the skin of the hand, back, and breast using a punctate probe. The first of the two touches (the reference) was at the same position on each trial, and the second touch (the comparison) was either above or below the first at a pre-specified distance. For the hand, distances ranged from 1 to 10 mm; for the other body regions, distances ranged from 2.5 to 40 mm, anticipating lower acuity based on preliminary testing (Figure 1). The participant’s task was to report whether the test stimulus was located above or below the reference stimulus. We then assessed the participant’s performance as a function of the distance between the test and reference and characterized the distance from the reference required to reliably locate the test stimulus – the just noticeable difference (JND) – for each participant and region.

Location discrimination.

A|Example psychometric functions for one subject on the location discrimination task for the hand, back, outer breast, and medial breast (including NAC). Negative values denote test points “below” the reference and the JND indicates the distance at which the subject could reliably (75% trials) locate the stimulus. B| Distribution of JNDs for each subject at each region. C| Relationship between the difference in size between the bust and underbust (Δ Bust) and the JND of the lateral breast. D| Relationship between the JND for the hand and breast across participants.

As expected, the hand demonstrated significantly greater spatial acuity (median [25th,75th] percentiles: 2.16 [1.73, 2.93] mm) than either the back (5.72 [4.36, 8.16] mm) or either breast region (6.63 [5.41, 7.73] mm and 8.00 [5.72, 10.80] mm, for the lateral and medial breast respectively) as indicated by the lower JNDs (Figure 1A, B, 1-way ANOVA: F[3,126] = 22.7, p < 0.01; Tukey’s HSD post-hoc t-test: p < 0.01). Interestingly, the medial breast yielded even shallower functions and lower spatial acuity than did the back (Figure 1A, B, 1-way ANOVA with Tukey’s HSD post-hoc t-test: p = 0.0284), though no statistical differences were observed between the lateral breast and either the back or medial breast (p = 0.3974 and 0.7377, respectively). In other words, touches needed to be between 3 and 4 times as far apart on the breast than on the hand to yield equivalent location discrimination performance.

Tactile acuity is worse for women with larger breasts

The spatial acuity of the hand has been shown to depend on the size of the hand, with smaller hands yielding better acuity (5). With this observation in mind, we investigated whether the inter-participant differences in breast acuity might be driven in part by differences in breast size. To determine the relative expansion (increase in surface area) of each participant’s breast, we computed the difference between their bust (circumference of the torso at the nipple line) and underbust (circumference of the torso at the inframammary fold). Comparing this value with their spatial acuity, we found that the acuity of the lateral breast decreased (JNDs increased) as breast size increased (Figure 1C, Pearson’s correlation with Bonferroni PHC (n=4): r = 0.734, p < 0.01), but found no effect at any other region (Supplementary Figure 1A-C, p > 0.05 for all), implying that any expansion is limited to the skin of the lateral breast and that the observed relationship between breast size and acuity was not spurious.

Earlier version: The variance in spatial acuity across body locations was similar across participants (2-way ANOVA: location– F[3,115] = 26.9, p < 0.01, η2 = 37%; participant - F[28,115] = 1.9, p = 0.0116, η2 = 25%). To examine if participants demonstrated consistent relative acuity across regions, we first ranked the acuity performance of each participant within region and then took the average rank across regions to determine the overall acuity rank of each participant. Measuring the rank correlation between the overall rank and that of each region, we found that acuity of the hand, lateral, and medial breast were all significantly correlated with overall rank (Kendall’s correlation with Bonferroni PHC: p < 0.01 for all) but that back acuity was not (p = 0.13 after PHC). While this observation implies that the relative skill or attention of each participant influenced the acuity of each region, differences in the spatial expansion of a region could likely override this effect. Thus, we also examined if the acuity of each region was predictive of any other by comparing their JNDs. Consistent with the fixed innervation hypothesis, we broadly found no relationship between the spatial acuity of each region, however, we found that the JNDs of the hand and medial breast were strongly related (Figure 1D, Pearson’s correlation, r = 0.685, p < 0.01; r = -0.007 to 0.375, p = 0.25 – 0.99 for other comparisons, Supplementary Figure 1D-H; Bonferroni PHC used for all where n = 6). Consequently, the spatial expansion that occurs at the medial breast is likely similar to that of the hand, implying that expansion of the NAC may be independent of the remainder of the breast.

Revised version: To examine if participants demonstrated consistent relative acuity across regions, we first ranked the acuity performance of each participant within region and then took the average rank across regions to determine the overall acuity rank of each participant. Measuring the rank-correlation between the overall rank and that of each region, we found that acuity of the hand, lateral, and medial breast were all significantly correlated with overall rank (Kendall’s correlation with Bonferroni PHC: p < 0.01 for all) but that back acuity was not (p = 0.13 after PHC). While this observation implies that the relative skill or attention of each participant influenced the acuity of each region, differences in the spatial expansion of a region could likely override this effect. Thus, we also examined if the acuity of each region was predictive of any other by comparing their JNDs. Consistent with the fixed innervation hypothesis, we found no relationship between the spatial acuity of each region, however, we found that the JNDs of the hand and medial breast were strongly related (Figure 1D, Pearson’s correlation, r = 0.685, p < 0.01; r = −0.007 to 0.375, p = 0.25 – 0.99 for other comparisons, Supplementary Figure 1D-H; Bonferroni PHC used for all where n = 6), though no efforts were made to measure hand size to examine if the two were related. Consequently, the spatial expansion that occurs at the medial breast is likely similar to that of the hand, implying that expansion of the NAC may be independent of the remainder of the breast.

The nipple is a unit

Earlier version: While the result that the medial breast has the lowest spatial acuity is not necessarily surprising given the lack of mechanoreceptors near the surface of the skin (13), the magnitude of the correlation was unexpected. Indeed, we observed that the JNDs for each participant were, on average, only marginally smaller than the measured diameter of the nipple (80% ± 52%). This would imply that 2 points on opposite halves of the nipple (80% diameter apart) might be confused as the same point. Given that the medial breast is inherently a heterogeneous structure as it is composed of three anatomically distinct regions of the breast (the nipple, the areola, and the nearby outer breast), we reasoned that this approach may be biased and thus not reflect the functional ability to locate sensations on the breast.

Revised version: While the result that the medial breast has the lowest spatial acuity is not necessarily surprising given the lack of mechanoreceptors near the surface of the skin (13), the magnitude of the correlation was unexpected. Indeed, we observed that the JNDs for each participant were, on average, only marginally smaller than the measured diameter of the nipple (80% ± 52%). This would imply that 2 points on opposite halves of the nipple (80% diameter apart) might be confused as the same point. Given that the NAC is inherently a heterogeneous structure as it is composed of three anatomically distinct regions of the breast (the nipple, the areola, and the nearby outer breast), we reasoned that this approach may be biased and thus not reflect the functional ability to locate sensations on the breast.

Earlier version: To assess if participants could reliably perceive tactile events across the nipple-areola complex, we delivered punctate touches to each quadrant on the nipple or areola and the participant reported the quadrant in which the touch had been delivered (Figure 2A). In particular, we marked the edge of the nipple or the 50% point between the center of the nipple and the edge of the areola in each of the cardinal directions such that the proportional distance was constant across participants. When touching different quadrants of the areola, participants were able to reliably locate stimuli (Figure 2B, mean ± standard deviation: 55% ± 21% where chance is 25%, one-sided Monte-Carlo test, Z = 15, p < 0.01), with 8 of the 10 participants performing significantly above chance (p < 0.05 after Bonferroni PHC where n = 10). Of the remaining participants, one performed marginally above chance (40%, p = 0.25 after Bonferroni PHC) while the other systematically misperceived the location and performed well below chance (7.5%, p = 1 after Bonferroni PHC). On the nipple, however, participants could only barely locate stimuli above chance (Figure 2B, 36% ± 13%; Z = 5.5, p < 0.01). Indeed, they were consistently worse at locating stimuli on the nipple than the breast (paired t-test, t = 3.42, p < 0.01) where only 3 of the 10 participants outperformed chance when considered individually. This finding suggests that the ability to locate tactile events is significantly worse at the nipple than at the areola and re-affirms our earlier finding that spatial acuity at the medial breast is poor such that contact events to it can only be poorly localized at best. Finally, to determine if this task was also influenced by spatial expansion, we sought to compare the breast size with performance on the quadrant task. Although we found negative trends for performance at both the areola and nipple, the data was collected post-hoc and only 4 of the 10 participants returned for measurements, thereby precluding statistical analysis (Figure 2C).

Revised version: To assess if participants could reliably perceive tactile events across the nipple-areola complex, we delivered punctate touches to each quadrant on the nipple or areola and the participant reported the quadrant in which the touch had been delivered (Figure 2A). In particular, we marked the edge of the nipple or the 50% point between the center of the nipple and the edge of the areola in each of the cardinal directions such that the proportional distance was constant across participants. When touching different quadrants of the areola, participants were able to reliably locate stimuli (Figure 2B, mean ± standard deviation: 55% ± 21% where chance is 25%, one-sided Monte-Carlo test, Z = 15, p < 0.01), with 8 of the 10 participants performing significantly above chance (p < 0.05 after Bonferroni PHC where n = 10). Of the remaining participants, one performed marginally above chance (40%, p = 0.25 after Bonferroni PHC) while the other systematically misperceived the location and performed well below chance (7.5%, p = 1 after Bonferroni PHC). On the nipple, however, participants were consistently worse at locating stimuli on the nipple than the breast (paired t-test, t = 3.42, p < 0.01) where only 3 of the 10 participants outperformed chance, though the group as a whole outperformed chance (Figure 2B, 36% ± 13%; Z = 5.5, p < 0.01). This finding suggests that the ability to locate tactile events is significantly worse at the nipple than at the areola and re-affirms our earlier finding that spatial acuity at the medial breast is poor such that contact events to it can only be poorly localized at best. Finally, to determine if this task was also influenced by spatial expansion, we sought to compare the breast size with performance on the quadrant task. Although we found negative trends for performance at both the areola and nipple, the data was collected post-hoc and only 4 of the 10 participants returned for measurements, thereby precluding statistical analysis (Figure 2C). This negative trend implicitly contradicts the previous result where one might expect equal performance regardless of size as the location of the stimuli was scaled to the size of the nipple and areola. However, given the absence of a reference point, systematic biases are more likely to occur and thus this may reflect a relationship between localization bias and breast size.

Spatial acuity of the nipple and areola.

A|Test locations for the quadrant discrimination task. The subject reported location using a number from 1 to 4 progressing clockwise from “above.” B| Proportion correct for participant at each test location. Solid markers indicate participants whose performance was significantly above chance (N=10). C| Relationship between breast size (bust – underbust) and performance on the quadrant task on the areola or nipple.

Earlier version: Absolute localization on the breast is anchored to the nipple

Revised version: Localization of tactile events on the breast is biased towards the nipple

Earlier version: In the experiments described above, we investigated participants’ ability to distinguish the relative locations of two touches to the breast. Next, we examined their ability to identify the absolute location of a tactile stimulus on their torso. To this end, we presented a single touch to the breast or the back with a punctate probe at one of 25 locations after which each participant marked the location of the touch on a three-dimensional digital image of the participant’s own breast or back (Figure 3A, B).

Revised version: In the experiments described above, we investigated the participants’ ability to distinguish the relative locations of two touches to the breast. Next, we examined their ability to identify the absolute location of a tactile stimulus on their torso. To this end, we presented a single touch to the breast or the back with a punctate probe at one of 25 locations after which each participant marked the location of the touch on a three-dimensional digital image of the participant’s own breast or back (Figure 3A, B).

Absolute localization of contact events on the breast and back.

A|Example localization task data for the breast or B| back of one participant. Gray surface represents scanned torso. Black crosses indicate the true location of each stimulus, purple lines indicate the vector between the stimulus location and reported location, and blue lines indicate the vector between the stimulus location and average reported location across blocks. C| Reporting error (3D Euclidean distance) as a function of distance across participants for the breast and D| back. Error indicates mean error for each trial (mean of purple vector length in A,B), bias indicates the error of the average response (blue vectors in A,B), and imprecision is the mean pairwise error between individual responses for each point (distance between the end of purple vectors in A,B). E| The distribution of angles (2D, no depth axis) between the stimulus location and the reported location (gray) or the difference between said angle and the angle towards the relevant landmark (red) for the breast (nipple) and F| back (scapula).

Earlier version: To estimate the degree to which participants could accurately localize touch events, we computed the error for each body part across participants and found that, contrary to the acuity task, participants had lower errors on their breasts in comparison to their back (Supplementary Figure 2A; paired t-test, t[9] = 7.4, p < 0.001). Next, we computed this error as a function of the distance from the center landmark for each region (nipple or scapula, Figure 3C, D) and found that in the case of the breast, the errors increased systematically as the distance from the nipple increased consistent with the quadrant localization task (2-way ANOVA: participant [9,44] F = 15.49, p < 0.001, distance [4, 44] F = 6.31, p < 0.001). This association was not found for the back (2-way ANOVA: participant [9,42] F = 0.79, p = 0.6315, distance [4, 42] F = 0.21, p = 0.9322). Given that any errors are likely a result of two sources of noise: bias and response variability (imprecision), we next sought to quantify the relative contribution of the two sources. We estimated the bias for each point by computing the centroid of all reports for a given stimulus across trials and measured the error between the centroid and the stimulus location. To measure the imprecision, we computed the mean pairwise error between each report and the stimulus location (Figure 3C, D). We found that, in both the breast and back, the error that originated from systematic bias significantly outweighed that of the imprecision (3-way ANOVA breast: participant [9,85] F = 16.12, p < 0.001, distance [4, 85] F = 5.6, p < 0.001, error type [1, 85] F = 17.4, p < 0.001; back: participant [9,85] F = 1.22, p = 0.2976, distance [4, 85] F = 0.14, p = 0.9655, error type [1, 85] F = 98.75, p < 0.001).

Revised version: To estimate the degree to which participants could accurately localize touch events, we computed the error for each body part across participants and found that, contrary to the acuity task, participants had lower errors on their breasts in comparison to their back (Supplementary Figure 2A; paired t-test, t[9] = 7.4, p < 0.001). Next, we computed this error as a function of the distance from the center landmark for each region (nipple or the inferior angle of the scapula, i.e. the bottom corner, Figure 3C, D) and found that in the case of the breast, the errors increased systematically as the distance from the nipple increased consistent with the quadrant localization task (Linear Mixed Model, LMM: distance: n = 248, B = 0.234, t = 9.194, p < 0.001, Table 1). This association was not found for the back (distance: n = 248, B = 0.093, t = 1.714, p = 0.087, Table 2). Then, using an LMM to compare the groups, we found both a significant interaction between the effect of distance and the location (location x distance: n = 496, B = −0.139, t = −2.723, p = 0.006, Table 3). Given that any errors are likely a result of two sources of noise: bias and response variability (imprecision), we next sought to quantify the relative contribution of the two sources. We estimated the bias for each point by computing the centroid of all reports for a given stimulus across trials and measured the error between the centroid and the stimulus location. To measure the imprecision, we computed the mean pairwise distance between each of the reported locations for a given stimulus location and the mean of those reported locations (Figure 3C, D). We found that, in both the breast and back, the error that originated from systematic bias significantly outweighed that of the imprecision (location x error type: n = 995, B = −23.05, t = 12.728, p < 0.001).

Earlier version: As bias was prevalent in the localization task, we next sought to determine if the bias was systematic across stimulus locations. Consequently, we assessed the tendency for the error vector to be uniformly distributed in two dimensions (Supplementary Figure 2B-C). Importantly, given the relative over-representation of the depth axis for the breast, we only considered angles in the horizontal planes such that the shape of the breast did not influence the results. First, we computed the absolute angle between stimulus location and the average reported location (blue vectors in Figure 3A, B) and computed the distribution across participants (Figure 3E, F). Across participants, we found no consistent excursions in the absolute biases from a normal distribution for the breast or back (permutation test for uniformity: R = 0.1314, p = 0.9331 and R = 0.1856, p = 0.7171, respectively).

Revised version: As bias was prevalent in the localization task, we next sought to determine if the bias was systematic across stimulus locations. Consequently, we assessed the tendency for the error vector to be uniformly distributed in two dimensions (Supplementary Figure 2B-C). Importantly, because the back is a relatively flat surface in comparison to the breast, error vectors were only computed in the horizontal plane and depth was excluded when computing the angular error. When computing the absolute angle between stimulus location and the average reported location (blue vectors in Figure 3A, B) and computed the distribution across participants (Figure 3E, F). Across participants, we found no consistent excursions in the absolute biases from a normal distribution for the breast or back (permutation test for uniformity: R = 0.1314, p = 0.9331 and R = 0.1856, p = 0.7171, respectively). Next, when we assessed if the bias with respect to the center point (nipple or scapula) was consistent, we found that both the breast and back tended to exhibit biases (R = 0.3140, p = 0.0100 and R = 0.6164, p < 0.0001, respectively), with the effect roughly doubled for the breast.

Earlier version: When we assessed if the bias with respect to the center point (nipple or scapula) was consistent, we found that both the breast and back tended to exhibit biases (permutation test for uniformity: R = 0.3140, p = 0.0100 and R = 0.6164, p < 0.0001, respectively), though the effect was roughly double for the breast. To further quantify this observation, we computed the vector strength of the biases for individual participants. For the breast we found that the biases tended to be uniformly distributed across participants, though some participants demonstrated biases (e.g. towards the top of the shoulder, Supplementary Figure 2D-F; vector strength Monte-Carlo test: p < 0.05 for 4/10 participants). Examining the back, there was a significant tendency for all but 1 participant to be biased towards a given direction, however, the direction was inconsistent across participants implying that reports were not biased towards any single landmark. Next, we computed the distribution of relative angles between the bias vector and the center point of our task and found that 7 of 10 participants were biased towards the nipple (vector strength Monte-Carlo: p < 0.05) and 2 were biased towards the scapula, consistent with the notion that biases are significant but inconsistent. Consequently, we conclude that while the back acuity is broadly low and subject to large errors, breast acuity is higher near the nipple and that perception of contact events across the breast are biased as being closer to the nipple.

Revised version: To further quantify this observation, we computed the vector strength – a measure of circular uniformity - of the biases for individual participants. For the breast we found that the absolute angle of the biases tended to be uniformly distributed across participants, though some participants demonstrated biases (e.g. towards the top of the shoulder, Supplementary Figure 2D-F; vector strength Monte-Carlo test: p < 0.05 for 4/10 participants). Examining the back, there was a significant tendency for all but 1 participant to be biased towards a given direction, however, the direction was inconsistent across participants implying that reports were not biased towards any single landmark. Next, we computed the distribution of relative angles between the bias vector and the center point of our task and found that 7 of 10 participants were biased towards the nipple (vector strength Monte-Carlo: p < 0.05) and 2 were biased towards the scapula, consistent with the notion that biases are significant but inconsistent.

Added in the revised version: Given these findings, we conclude that the breast has lower tactile acuity than the hand and is instead comparable to the back. Moreover, localization of tactile events to both the back and breast are inaccurate but localizations to the breast are consistently biased towards the nipple.

Discussion

First, we found that the spatial acuity of the lateral breast – excluding the nipple and areola – is almost four times lower than that of the hand, and even lower than that of the back, previously considered the epitome of poor acuity. Second, the nipple is essentially a sensory unit, as touches to different aspects of the nipple are nearly indistinguishable from one another. Third, the acuity of the breast tends to be poorer for women with large breasts, consistent with the theory that innervation capacity is fixed. Together, the data indicates that the nipple constitutes a landmark on the breast, as evidenced by the fact that absolute localization judgments are less accurate and more biased for touches that are far from the nipple and perceived location is pulled systematically toward the nipple.

The poor spatial acuity of the breast

Earlier version: The poor spatial acuity of the breast – about four times lower than that of the hand and slightly worse than the back – replicates previous findings (8), though these findings were achieved using less reliable methods – two point threshold and a ‘same-different’ paradigm. Tactile spatial acuity is determined by density of innervation: more neural tissue is devoted to more highly innervated body regions and this increased central representation is a key contributor to the increased acuity. The low spatial acuity of the nipple and areola is consistent with a histological study revealing these regions to be sparsely innervated (13).

Revised version: The poor spatial acuity of the breast – about four times lower than that of the hand and slightly worse than the back – replicates previous findings (8) achieved using less reliable methods (two point threshold and a ‘same-different’ paradigm). Tactile spatial acuity is determined by density of innervation: more neural tissue is devoted to more highly innervated body regions and this increased central representation is a key contributor to the increased acuity. The low spatial acuity of the nipple and areola is consistent with a histological study revealing these regions to be sparsely innervated (13).

Larger breasts confer lower acuity

Earlier version: We found a significant relationship between breast size and spatial acuity: women with larger breasts tended to exhibit lower spatial acuity on their breasts, consistent with previous findings that tactile acuity scales with body size. Indeed, the spatial acuity of the hand has been shown to depend on the size of the hand, with smaller hands exhibiting better acuity (5, 6). These results are consistent with the hypothesis that the number of tactile nerve fibers does not scale with body size, so fibers are more sparsely distributed on bigger bodies, leading to lower acuity.

Revised version: We found a significant relationship between breast size and spatial acuity: women with larger breasts tended to exhibit lower spatial acuity on their breasts, consistent with previous findings that tactile acuity scales with body size. Indeed, the spatial acuity of the hand has been shown to depend on the size of the hand, with smaller hands exhibiting better acuity (5, 6). These results are consistent with the hypothesis that the number of tactile nerve fibers does not scale with body size, so fibers are more sparsely distributed on bigger bodies, leading to lower acuity. Consequently, the acuity of the breast is likely determined initially by torso acuity and then any expansion.

Added in the revised version: Reports are biased towards the nipple

Removed in the revised version: The nipple is a landmark

Earlier version: The mental representation of the body is not uniform and veridical (14). Localization tends to be more precise when stimuli are applied near anatomical points of reference that form perceptual anchor points (15). For example, the navel and spine act as anchor points along the abdomen (16, 17): touches to the navel or spine are never mistaken for touches anywhere else on the abdomen. Furthermore, touches to locations near these two anchor points are mislocalized following a bias toward these areas: touches near the navel are pulled toward the navel and touches near the spine are pulled toward the spine. Similar biases are observed on the back of the hand, where touches are mislocalized to be closer to the fingers than they actually are (18), and on the forearm, where they are pulled toward the wrist (19). Analogously, we found that acuity is highest near the nipple, and perceived locations are pulled toward the nipple, consistent with the nipple acting as an anchor point on the breast.

Revised version: The mental representation of the body is not uniform and veridical (14). Localization tends to be more precise when stimuli are applied near anatomical points of reference that form perceptual anchor points (15). For example, the navel and spine act as anchor points along the abdomen (16, 17): touches to the navel or spine are never mistaken for touches anywhere else on the abdomen. Furthermore, touches to locations near these two anchor points are mislocalized following a bias toward these areas: touches near the navel are pulled toward the navel and touches near the spine are pulled toward the spine. Similar biases are observed on the back of the hand, where touches are mislocalized to be closer to the fingers than they actually are (18), and on the forearm, where they are pulled toward the wrist (19). Analogously, we found that acuity is highest near the nipple, and perceived locations are pulled toward the nipple, suggesting that the nipple plays a pivotal role in the mental representation of the breast.

Earlier version: Implications for a bionic breast

Revised version: Implications for breast prostheses

Removed in the revised version: Understanding the spatial acuity of the breast is particularly important given a recent proliferation of efforts - including using autologous tissue, synthetic grafts and neuroprosthetic or bionic approaches - to restore sensation to the breast following mastectomy (20, 21). In particular, this work shows that the spatial resolution of an implantable sensor sheet should depend on the distance from the nipple.

Earlier version: Individual sensors should be placed in each quadrant of the nipple while subsequent sensors should be placed radially with increasing separation (0.5 to 2 cm inter-sensor spacing) up to a 5 cm radius at which point resolution would remain constant at approximately 2 cm. These values, however, depend on the spatial resolution of the stimulation technology which can vary significantly for both peripheral nerve stimulation (22) and intracortical microstimulation (23).

Revised version: Understanding the spatial acuity of the breast is particularly important given a recent proliferation of efforts - including using autologous tissue, synthetic grafts and neuroprosthetic or bionic approaches – to restore sensation to the breast following mastectomy (20, 21). In particular, this work shows that the spatial resolution of an implantable sensor sheet should depend on the distance from the nipple. Individual sensors should be placed in each quadrant of the nipple-areolar complex while subsequent sensors should be placed radially with increasing separation (0.5 to 2 cm inter-sensor spacing) up to a 5 cm radius at which point resolution would remain constant at approximately 2 cm. These values, however, depend on the spatial resolution of the stimulation technology which can vary significantly for both peripheral nerve stimulation (22) and intracortical microstimulation (23).

Conclusion

The breast has unique sensory properties because (1) it mediates lactation and nursing, (2) it comprises distinct regions – the nipple, areola, and outer breast – which differ in the type of skin and patterns of sensory innervation, (3) it has erectile function and gives rise to erogenous sensations, and (4) it undergoes variable expansion across individuals during puberty. First, we find that spatial acuity on the breast is lower than that of the back, previously regarded as the body region with lowest tactile acuity. Second, spatial acuity is lower in larger breasts, presumably due to sensory innervation being fixed prior to expansion during puberty. Third, the nipple itself is perceived as a single spatial unit and serves as an anchor for the mental map of the breast.

Methods

Participants

A total of 48 healthy adult women (24.2 ± 3.1, 19-32 years) participated in this study: Thirty-four in the location discrimination task (mean ± standard deviation, age range; 22.6 ± 2.7, 19-28 years), ten in the medial breast quadrant localization task (25.3 ± 3.8, 19-32 years), and ten in the absolute localization task (24.7 ± 2.9, 20-28 years). Some participants performed multiple tasks. Experimental procedures were performed in accordance with the relevant guidelines and regulations and were approved by the Institutional Review Board of the University of Chicago (IRB 18-0135). The informed consent process included documentation of written consent from each participant, and the participants were compensated for their participation. We excluded any women who, by self-report, were currently pregnant or breastfeeding, had a history of breast surgery, or any diagnosis of neurological illness.

Breast measurements

In addition to self-reported bra size, we collected standardized, objective measurements from each participant. Specifically, we measured the bust – i.e. the circumference of the chest at the level of the nipple – and subtracted from it the under-bust, the circumference at the level of the inframammary fold (mean ± standard deviation, range; 12.4 ± 4.6, 7.6-25.4 cm). We also measured the diameter of the areola and nipple (areola: 37.4 ± 11.9, 21-80 mm; nipple: 12.7 ± 2.7, 7-20 mm).

Tactile Stimuli

Earlier version: Touches were delivered manually with a 1-mm tipped XP-Pen (XP-Pen USA, CA, USA) which allowed the experimenter to lightly pressed the stimulus against the skin and ensure that the skin was indented uniformly across touches. The XP-Pen was selected instead of monofilaments because, in pilot experiments, several participants reported discomfort with the sharp edges of the monofilaments on their breasts while also allowing for the application of a constant force. The participants were asked to report any discomfort with the application of the stimulus or the inability to feel the application of the stimulus reliably. Note that spatial acuity is consistent across stimulus amplitudes as long as the touch is sufficiently above threshold (24), which was the case here.

Revised version: Touches were delivered manually with a 1-mm tipped XP-Pen (XP-Pen USA, CA, USA) which allowed the experimenter to lightly press the stimulus against the skin and ensure that the skin was indented uniformly across touches. The XP-Pen was selected instead of monofilaments because, in pilot experiments, several participants reported discomfort with the sharp edges of the monofilaments on their breasts while also allowing for the application of a constant force. The participants were asked to report any discomfort with the application of the stimulus or the inability to feel the application of the stimulus reliably. Note that spatial acuity is consistent across stimulus amplitudes as long as the touch is sufficiently above threshold (24), which was the case here.

Psychophysical tasks

Location discrimination

Acuity was tested at each of four body regions: the lateral breast, the medial breast (which includes the NAC), the thenar eminence of the hand, and the upper back. The participant – wearing a gown that exposed only the location to be tested – lay supine on a massage table for testing on the lateral breast, medial breast, and thenar eminence, and prone for the testing on the back. Unless the participant expressed a preference, the side of the body to be tested (left/right) was chosen randomly (through a random number generator) but each participant was tested on the same side for all regions. The order in which the different regions were tested was randomized to eliminate any effects of fatigue or learning.

On each trial, two touches were applied to nearby locations and the participant’s task was to indicate whether the second touch was above or below the first by pressing one of two buttons on a keypad. The location of one of the two touches (the reference) was consistent across each experimental block and the location of the second (the comparison) varied from trial to trial (each at a pre-specified distance). Fifteen comparison locations, aligned along an axis parallel to the body’s axis, were tested on the back, lateral breast, and medial breast. Touch locations were drawn on the body to ensure repeatable presentation and comparisons were placed at locations 0, 2.5, 5, 7.5, 10, 15, 20 or 40 mm above or below the reference (Figure 1A). Each comparison was repeated five times per block over three blocks (for a total of 15 repeats per comparison, with 225 total trials per body region) in a randomized order. Thirteen comparison locations were tested on the hand, 0, 1, 3, 6, 7.5, and 10 mm away from the reference in both directions along an axis in line with the thumb For the hand, a judgment of “above” indicated that the comparison was displaced toward the thumb relative to the reference. Each location was repeated 5 times across three experimental blocks (for a total of 15 repeats per comparison, with 195 total trials).

Performance was gauged by the proportion of times the participant judged a comparison as above the reference as a function of distance from the reference, where negative distances indicate comparison locations below the reference. Psychometric functions were fit to the data and used to compute the just noticeable difference.

Quadrant discrimination

The objective of this experiment was to assess the degree to which women can distinguish touches to different parts of their NAC. On each trial, the participant was touched at one of four locations on the nipple or areola (organized in a quadrant) and verbally identified the touch location using a number from 1 to 4. Each location was touched ten times per block (40 trials per block, with 80 total trials per breast region). For the areola, each touch was located halfway between the edge of the areola and the nipple. For the nipple, each touch was delivered at the edge of the nipple, in line with each point on the areola (Figure 2A). The testing side was chosen randomly for each participant (6 left/4 right).

Absolute localization

Earlier version: The objective of this experiment was to gauge the accuracy with which women could report where on their breast or back a touch was delivered. On the breast, touches were arranged such that the nipple was the central point and other touches radiated outwards from it (Figure 3A). On the back, touches were arranged the same way around a central point located on the shoulder blade (Figure 3B). First, the participant’s skin was marked, then a three-dimensional scan of the breast or back was obtained using the EM3D application (Brawny Lads Software, LLC.) and .obj files were uploaded to Blender. The surface of the breast was then masked with a uniform layer of grey to obscure the markings. A laptop (13” screen) was positioned such that the participant could use the trackpad to interact with the 3D image (zoom, rotate, drag object) to report where the touch was experienced on the 3D rendering of their breast. Each of the 25 stimulus locations was presented in a pseudorandomized order while the participant closed their eyes. After the touch, the participant opened their eyes and marked the perceived location of the touch on the 3D rendering of their breast or back. The participant was encouraged to manipulate the 3D model freely to obtain the best view of their breast or back on each trial. The participant moved the cursor and clicked to indicate where they perceived the touch. Reported locations accumulated on the 3D model throughout the block (during which each location was touched once.) but were removed at the start of each of three blocks (yielding a total of 3 repeats for each of the 25 locations).

Revised version: The objective of this experiment was to gauge the accuracy with which women could report where on their breast or back a touch was delivered. On the breast, touches were arranged such that the nipple was the central point and other touches radiated outwards from it (Figure 3A). On the back, touches were arranged the same way around a central point located on the shoulder blade (Figure 3B). First, the participant’s skin was marked, then a three-dimensional scan of the breast or back was obtained using the EM3D application (Brawny Lads Software, LLC.) and .obj files were uploaded to Blender. The surface of the breast was then masked with a uniform layer of grey to obscure the markings. A laptop (13” screen) was positioned such that the participant could use the trackpad to interact with the 3D image (zoom, rotate, drag object) to report where the touch was experienced on the 3D rendering of their breast. Notably, in both cases landmarks were observable on the 3D models (e.g. nipple, shoulder, scapula, vertebral line). Each of the 25 stimulus locations was presented in a pseudorandomized order while the participants closed their eyes. After the touch, the participant opened their eyes and marked the perceived location of the touch on the 3D rendering of their breast or back. The participant was encouraged to manipulate the 3D model freely to obtain the best view of their breast or back on each trial. The participant moved the cursor and clicked to indicate where they perceived the touch. Reported locations accumulated on the 3D model throughout the block (during which each location was touched once.) to encourage placement of markers in accurate relative positions but were removed at the start of each of three blocks (yielding a total of 3 repeats for each of the 25 locations).

Data Analysis

Spatial Acuity

To quantify the spatial acuity at each test location, we fit a psychometric function (see EQ1) to the probability of judging the second touch as above the first versus the relative location of the two stimuli. From these functions, we estimated the distance from the reference point required to reliably discriminate the sensation (75% performance, see EQ2).

Where k is the growth term of the exponential function and x is the distance between the comparison point and the reference point.

Where k is the growth term computed in EQ1.

Added in the revised version: JNDs were then compared between regions using a 1-way ANOVA followed by a Tukey’s honest significant difference post-hoc test.

Absolute localization task error and bias

To compute the error for the task, we considered the 3D position of both the stimuli and response. First, to compute the total error, we computed the 3D Euclidean distance between the stimulus location and the participants response on each trial and then averaged these values together. While the skin is not flat, over the relatively short distances between stimulus and response (typically <5 cm), Euclidean distance was sufficient as minimal curvature occurred over this range. To measure the bias for each point, the 3D position for all responses for a given point across blocks was averaged together and the Euclidean distance between the resultant point and the stimulus location was averaged (centroid). If responses were randomly distributed around the stimulus location, then the error of the average response would be minimal. If there was a consistent offset, then the bias error would be similar to the total error. Finally, to compute imprecision, we computed the mean Euclidean error between the centroid and individual responses.

Added in the revised version: Linear mixed models were used to assess significant effects using Matlab’s fitlme function. The formula used for each test can be found in Tables 1–4. While the fixed effects varied, participants were treated as random effects and the intercept and slope were allowed to vary for each.

Removed in the revised version: When combining across participants, we had to account for the fact that there were subtle variations between participants in the absolute distances between the center point and each stimulus location due to skin stretch or torso morphology. Consequently, we discretized the stimuli by binning them within their groups to facilitate N-way ANOVAs. Thus, distances were binned between the following edges: 0, 5, 25, 40, 60, and 90 mm from the center.

Absolute localization task angular error

To compute the angle of the errors for each response we measured the four-quadrant inverse tangent (EQ3) between each point and response. Importantly, this was only considered in 2-dimensions as the variance of angles in the 3rd dimension (depth) were drastically dependent on breast size and the location of the point on the breast. Thus, the angle between the stimulus location and each response (as well as the centroid across responses) was computed. Then, for each stimulus the angle between the stimulus location and the reference point was computed (center angle) the difference in angle between said angle and the centroid angle was measured.

Absolute localization task uniformity test

Earlier version: To test if the cross-participant biases were uniform or not we measured the range-normalized mean squared error (rnMSE, EQ4) between the average distribution of error angles across participants and the mean of the distribution, similar to measurements of residual variance. Then, to determine if the observed rnMSE was significant, we used a permutation test in which for each permutation (n = 1e5) we shuffled the angles for each participant before averaging across participants and computing the rnMSE to produce a null distribution of rnMSEs that might be expected if the participants responses were not systematic. Finally, we computed the proportion of simulations that had greater rnMSEs than the observed rnMSE to produce the p value.

Revised version: To test if the cross-participant biases were uniform or not we measured the range-normalized mean squared error (rnMSE, EQ4) between the average distribution of error angles across participants and the mean of the distribution, similar to measurements of residual variance. Then, to determine if the observed rnMSE was significant, we used a permutation test in which for each permutation (n = 1e5) we shuffled the angles for each participant before averaging across participants and computing the rnMSE to produce a null distribution of rnMSEs that might be expected if the participants responses were not systematic. Finally, we computed the proportion of simulations that had greater rnMSEs than the observed rnMSE to produce the p value which was considered significant if it was less than 0.05.

Absolute localization task vector strength

The vector strength for individual participants was computed using EQ5. Importantly, given the observed heteroskedasticity between distance and error, we did not incorporate the length of the vectors in the computation as this would have biased the computation towards points distal to the nipple for the breast but would not have influenced the back measurements. As the angular errors were distributed between −π and π, to test if the observed vector strength was significant, we used a Monte-Carlo simulation in which for each simulation (n = 1e5) we sampled angles from the range [-π, π] from a uniform distribution (n = 11) and computed the vector strength from each simulation. We then compared the observed vector strengths for each participant and location against the null distribution and considered it significant if the value exceeded the 95th percentile of the null distribution.